We also want to look a the distribution of the `price` variable. Let's do that with a new type of graph: the box plot.

Box plots (also called box-and-whisker plots) visualize not only the distribution but also show several summary statistics:

TODO: replace with my own image that I know I own copyright of.

```{r}
set.seed(42)
x <- c(rnorm(100, 20, 3), rnorm(5, 20, 6))
arrow_df <- tibble(
  x_start = c(median(x), quantile(x)[2]-1.1, quantile(x)[4]+1.1, 12,        30),
  x_end   = c(median(x), quantile(x)[2],     quantile(x)[4],     11.1,      30.9),
  y_start = c(1.7,       1.7,                1.7,                1.5,       1.5),
  y_end   = c(1.45,      1.45,               1.45,               1.1,       1.1),
  label   = c("Median", "25th Percentile",   "75th Percentile",  "Outlier", "Outlier")
)
ggplot() +
  geom_boxplot(aes(
    x, 
    ""
    ),
    fill = "lightblue") +
  labs(
    x = "x-axis = range of plotted variable", 
    y = "",
    title = "Parts of a box plot") +
  theme_minimal() +
  geom_segment(
    aes(x = x_start, y = y_start, xend = x_end, yend = y_end),
    data = arrow_df, 
    arrow = arrow(length = unit(0.2, "cm"), type="closed")
    ) +
  geom_text(
    aes(x = x_start, y = y_start + 0.05, label = label),
    data = arrow_df
  ) +
  geom_brace( # upper whisker
    aes(
      c(
        quantile(x)[4]+0.1,
        27.1 #1.5*IQR(x)+quantile(x)[4]
        ), 
      c(0.59, 0.5)), 
    rotate = 180,
    inherit.data=FALSE
    ) +
  geom_brace( # lower whisker
    aes(
      c(
        quantile(x)[2] - 0.1,
        quantile(x)[2] - 1.5*IQR(x) +0.2
        ), 
      c(0.59, 0.5)), 
    rotate = 180,
    inherit.data=FALSE
    ) +
  geom_segment( # This arrow forces the y axis to extend high enough
    aes(x = 11, y = 1.9, xend = 11, yend = 1.45), 
    color = "white",
    arrow = arrow(length = unit(0.2, "cm"), type="closed")
    ) +
  annotate("text", x = 24.5, y = 0.45, label = "Upper \"whisker\"") +
  annotate("text", x = 15.0, y = 0.45, label = "Lower \"whisker\"")
```


```{r}
knitr::include_graphics("../img/boxplot.png")
```

We create a box plot using the `geom_boxplot` geom function. There are two required aesthetic mappings in a box plot (note that in `ggplot` the orientation of the boxplot is rotated 90 degrees): 

* `x` should be a categorical variable to create different box plots for.
* `y` is the variable that we want to show the distribution of.

```{r}
diamonds %>%
  ggplot() +
  geom_boxplot(mapping = aes(x = price))
```

One slightly strange feature of the default box plot is that `ggplot` will put some meaningless numbers on the y-axis. We can get rid of these by mapping the `y` aesthetic to an empty character string:

```{r}
ggplot(data = diamonds) +
  geom_boxplot(mapping = aes(x = price, y = ""))
```

This still leaves us with an unnecessary y axis label, but we can also get rid of that with the `labs()` function.

<!-- Book version -->

TODO: book exercises

<!-- End book -->

<!-- Tutorial version -->

<p class="comment3">
Create a single box plot of the `price` variable by providing the string `"price"` as the argument to the `x` parameter.
</p>

```{r ixzvzwepjdibpkvq, exercise = TRUE, eval=is_tutorial, echo=FALSE}
# Fill in the blank:
diamonds %>%
  ggplot() +
  geom_boxplot(mapping = aes(x = _____, y = price))
```

```{r ixzvzwepjdibpkvq-solution, eval=is_tutorial, echo=FALSE}
# solution code
diamonds %>%
  ggplot() +
  geom_boxplot(mapping = aes(x = "price", y = price))
```

```{r ixzvzwepjdibpkvq-check, eval=is_tutorial, echo=FALSE}
# check code
gradethis::grade_code()
```

<!-- End tutorial -->
